---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Global and local outliers (version 1)

```yaml
type: VideoExercise
key: c69cd5603d
xp: 50
```

`@projector_key`
8ea34aee42f724715fb4e127061371c3

---

## Global and Local Outlier (version 2 with code as requested)

```yaml
type: VideoExercise
key: a0ce175a59
xp: 50
```

`@projector_key`
7f271ff2377d276f54a24c31cffa28a7
